# Homework 05

In this exercise, you will practice inferential statistics with confidence intervals, bootstrapping, and hypothesis testing. Problems may involve a combination of math and code. 

Recall that you can use LaTeX to nicely format your math inside Markdown cellsby enclosing equations in single dollar signs (e.g., $x^2+4=8$) for inline math or double dollar signs for centered equations like $$P(X > 5) = \frac{1}{6}.$$ See Worked Examples/HW 03 for further instructions and more examples.

Show your work and/or briefly explain your answers. In general, you will not receive full credit for numeric answers with no accompanying work or justification (math, code, explanation). For numeric answers, we will accept answers that are very slightly off due to rounding, $z$-score of $2$ vs. $1.96$, etc. 

```python
# Run this code cell to import relevant libraries
import numpy as np
import pandas as pd
from scipy import stats
```

### Question 1 (4 points, all autograded)

A website is trying to increase registration for first-time visitors, exposing a random subset of these visitors to a new site design. Of $752$ randomly sampled visitors over a month who saw the new design, $64$ registered. Using [`stats.norm.interval`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html), construct a $95\%$ confidence interval for the percentage of visitors who would register for the website under the new design. 

Note that the first parameter is named `confidence` in `scipy` 1.9.0+ and `alpha` in `scipy` 1.8.1- in both functions. **The autograder uses `scipy` 1.9.0+, so specifying `alpha` will cause your cell to error out on Gradescope. You can either use `confidence` (if you have `scipy` 1.9.0+), or just not specify the first parameter name, e.g., `stats.norm.interval(0.95, ...)`.**

Save your answer in a tuple `q1_1` of two `float` or `numpy.float64` variables such that `q1_1[0]` is the lower/left bound and `q1_1[1]` is the upper/right bound of your confidence interval. We accept either percentages (e.g., $50.0$) or fractions (e.g., $0.5$). [`stats.norm.interval`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html) when used correctly already returns a tuple.

```python
# Code for question 1
web_scale = np.sqrt((64/752)*(1-64/752))/(np.sqrt(752))
web_interval = stats.norm.interval(0.95, loc = (64/752), scale = web_scale)
q1 = web_interval


# Leave these lines here for grading and ease of debugging
print("Q1: 95% confidence interval: ", q1)
```

Output:
```
Q1: 95% confidence interval:  (0.06516269200219607, 0.10505007395525073)
```

### Question 2 (8 points: 4 autograded, 4 manually graded)

A study examined the mean pay for a random sample of men and women entering the workforce as doctors for $21$ different positions. If each gender was equally paid, then we would expect about half of those positions to have men paid more than women and women would be paid more than men in the other half of positions. In the study, men were, on average, paid more in $17$ of the $21$ positions. 

Complete a hypothesis test using `stats.norm.cdf` to examine whether there is significant evidence (at the $0.05$ level) of gender discrimination in pay in these positions. It can be either one-sided or two-sided. Uncomment the corresponding line to indicate which test you are conducting in the variable `q2_setting`. Save your p-value in `q2` as `float` or `numpy.float64`. 

In the written part:
  1. Explicitly write down the null hypothesis and the alternative hypothesis for this test.
  1. Report your p-value and interpret the result.
  1. Explain why using `stats.norm.cdf` (not `stats.t.cdf`) is a good choice for this context.

```python
# Code for question 2
# q2_setting = 1 # Uncomment for one-sided test
q2_setting = 2 # Uncomment for two-sided test

job_z_score = np.sqrt(21)*((17/21)-0.5)/(np.sqrt(0.5*(1-0.5)))

q2 = (1-stats.norm.cdf(job_z_score))*2

# Leave these lines here for grading and ease of debugging
print("Q2: p-value = {} for {}-sided test".format(q2, q2_setting))
```

Output:
```
Q2: p-value = 0.004556349803185089 for 2-sided test
```

<!-- BEGIN QUESTION -->

### Answer 2

For a two-sided test, let the null hypothesis $H_0$ be that $50\%$ of positions have men paid more than women. The alternative hypothesis $H_a$ be that something other than $50\%$ of positions have men being paid more than women.

Conducting a hypothesis test using the normal distribution, I found a two-sided p-value of around $0.005$ for the probability that we observe something as extreme as $50\%$ given the $H_0$. This is statistically significant at the $0.05$ level because the p-value is less than $0.05$, so we can reject the null hypothesis. Rejecting the null hypothesis shows there is gender discrimination in pay of job positions, as there is not an even amount of which gender is being paid more in job positions. 

stats.norm.cdf is a good choice because the data is proportional, so we can find the mean and standard deviation in order to calculate a z-score to use the normal distribution.

<!-- END QUESTION -->

## Movie Ratings Data
In the remainder of this assignment you will work with the movielens dataset of movie ratings that we have seen before. Below we import and preview the data. It consists of 2 tables: `users` has a row for every individual who has rated any movies, `movie-ratings` has a row for every rating of a particular movie by a particular user. This means users with multiple ratings are in the `movie_ratings` multiple times. The data is a random sample of all of the movie ratings made on the movielens service.

```python
users = pd.read_csv("users.csv")
users.head()
```

Output:
```
user_id  age sex  occupation
0        1   24   M  technician
1        2   53   F       other
2        3   23   M      writer
3        4   24   M  technician
4        5   33   F       other
```

```python
movie_ratings = pd.read_csv("movies-all.csv")
movie_ratings.head()
```

Output:
```
user_id  age sex  occupation  movie_id  rating                 movie_title
0        1   24   M  technician        61       4  Three Colors: White (1994)
1       13   47   M    educator        61       4  Three Colors: White (1994)
2       18   35   F       other        61       4  Three Colors: White (1994)
3       58   27   M  programmer        61       5  Three Colors: White (1994)
4       59   49   M    educator        61       4  Three Colors: White (1994)
```

### Question 3 (12 points: 8 autograded, 4 manually graded)
1. Compute a $95\%$ confidence interval for the mean `age` of users using `stats.norm.interval`. Save your answer in a tuple `q3_1` similar to that in Q1.
2. Compute a $95\%$ confidence interval for the mean `age` of users **who have rated the movie `Casablanca (1942)`** using the normal distribution. Save your answer in a tuple `q3_2`.

In the written part, answer the following question.

3. `Casablanca (1942)` is an old movie, one might suspect that it has been rated by older individuals on average than the entire dataset. Just looking at the confidence intervals you computed in steps 1 and 2, can you conclude that there is significant evidence for this belief? Why or why not? 

```python
# Code for question 3
users_mu = np.mean(users['age'])
users_sigma = np.std(users['age'])
users_n = len(users['age'])
q3_1 = stats.norm.interval(0.95, users_mu, users_sigma/np.sqrt(users_n))

casablanca = movie_ratings[movie_ratings['movie_title'] == 'Casablanca (1942)']
casablanca_mu = np.mean(casablanca['age'])
casablanca_sigma = np.std(casablanca['age'])
casablanca_n = len(casablanca['age'])
q3_2 = stats.norm.interval(0.95, casablanca_mu, casablanca_sigma/np.sqrt(casablanca_n))

# Leave these lines here for grading and ease of debugging
print("3.1: 95% confidence interval: ", q3_1)
print("3.2: 95% confidence interval: ", q3_2)
```

Output:
```
3.1: 95% confidence interval:  (33.27417039488504, 34.829753253047095)
3.2: 95% confidence interval:  (34.46345637868268, 37.330782304444895)
```

<!-- BEGIN QUESTION -->

## Answer 3

Since there is overlap between the two confidence intervals, we can't conclude there is significant evidence for the belief that Casablanca is rated by older individuals on average than the entire dataset. In order to determine if there was significance for this belief, we would have to perform a t-test.

<!-- END QUESTION -->

### Question 4 (6 points, all autograded)
Only $18$ users have rated the movie `Lost in Space (1998)`.
1. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the mean `age` of users who have rated `Lost in Space (1998)`. Save your answer in a tuple `q4_1`.
2. One of the advantages of bootstrapping is that we can easily compute confidence intervals for arbitrary measurements of distributions. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the **median** `rating` of `Lost in Space (1998)`. Note that `numpy` provides a vectorized function for [calculating the median](https://numpy.org/doc/stable/reference/generated/numpy.median.html) as well as the mean. Save your answer in a tuple `q4_2`.

```python
# Code for question 4
lost_in_space = movie_ratings[movie_ratings['movie_title'] == 'Lost in Space (1998)']
ages_lost = lost_in_space['age']
num_bootstraps = 10000

movie_bootstrap = np.random.choice(ages_lost, size=(num_bootstraps, len(ages_lost)), replace = True)
movie_bootstrap_means = np.mean(movie_bootstrap, axis=1)
movie_boot_l = np.percentile(movie_bootstrap_means, 2.5)
movie_boot_r = np.percentile(movie_bootstrap_means, 97.5)

q4_1 = (movie_boot_l, movie_boot_r)

rating_lost = lost_in_space['rating']
rating_bootstrap = np.random.choice(rating_lost, size=(num_bootstraps, len(rating_lost)), replace = True)
rating_bootstrap_median = np.median(rating_bootstrap, axis=1)
rating_boot_l = np.percentile(rating_bootstrap_median, 2.5)
rating_boot_r = np.percentile(rating_bootstrap_median, 97.5)

q4_2 = (rating_boot_l, rating_boot_r)

# Leave these lines here for grading and ease of debugging
print("4.1: 95% confidence interval: ", q4_1)
print("4.2: 95% confidence interval: ", q4_2)
```

Output:
```
4.1: 95% confidence interval:  (25.944444444444443, 36.611111111111114)
4.2: 95% confidence interval:  (2.5, 4.0)
```

### Question 5 (14 points: 6 autograded, 8 manually graded)
The `Star Wars (1977)` film is quite popular, with a median rating of $5/5$. However, of those that left a rating, male users gave it a slightly higher mean rating of about $4.4$ whereas female users gave the same movie a mean rating of about $4.2$.

1. Consider the null hypothesis that the mean rating of `Star Wars (1977)` is the same for `sex='F'` and `sex='M'` users. The alternative hypothesis is that the mean ratings are not equal. Conduct a two-sided t test using [`stats.ttest_ind`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind.html) to evaluate this using the sample ratings data. Report your p-value in `q5_1` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell. **Moreover, explain the underlying assumptions in your t test in the cell.** You are welcome to consult the documentation of [`stats.ttest_ind`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind.html) to help answering this question.


2. Consider the null hypothesis that $51\%$ of men would rate `Star Wars (1977)` a $5$. The alternative hypothesis is then that the fraction of men who would rate `Star Wars (1977)` a $5$ is not $51\%$. Conduct a two-sided hypothesis test using [`stats.t.cdf`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.t.html) to evaluate this in light of the sample ratings data of male users who rated `Star Wars (1977)`. Report your p-value in `q5_2` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell.


3. Consider the null hypothesis that women and men were equally likely to rate `Star Wars (1977)` a $5$. The alternative hypothesis is that these proportions are not equal. Conduct a two-sided t test using [`stats.ttest_ind_from_stats`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind_from_stats.html) to evaluate this in light of the sample data of female and male users who rated `Star Wars (1977)`. Report your p-value in `q5_3` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell. 


4. In the sample ratings data, $51\%$ of women rated `Star Wars (1977)` a $5$. (You are welcome to write code to validate this claim, but you can just accept it as fact.) Therefore, if $51\%$ of men and $51\%$ of women would rate `Star Wars (1977)` a $5$, both null hypotheses in 5.2 and 5.3 would be true. If this is the case, why do you observe a different p-value in the two parts, despite the hypotheses under consideration being ostensibly similar? Briefly explain why you observe this difference in the **Answer 5** cell. 

```python
# Code for question 5
star_wars = movie_ratings[movie_ratings['movie_title'] == 'Star Wars (1977)']
women_rank = star_wars[star_wars['sex'] == 'F']['rating']
# print(len(women_rank))
men_rank = star_wars[star_wars['sex'] == 'M']['rating']
# print(len(men_rank))
q5_1 = stats.ttest_ind(women_rank, men_rank)[1]

sw_empirical = len(men_rank[men_rank == 5])/len(men_rank)
sw_t_score = np.sqrt(len(men_rank))*((sw_empirical)-0.51)/(np.sqrt(0.51*(1-0.51)))
q5_2 = (1-stats.t.cdf(sw_t_score, len(men_rank)-1))*2

men_empirical = len(men_rank[men_rank == 5])/len(men_rank)
women_empirical = len(women_rank[women_rank == 5])/len(women_rank)
q5_3 = stats.ttest_ind_from_stats(mean1 = men_empirical, std1 = np.sqrt(men_empirical*(1-men_empirical)), nobs1 = len(men_rank), 
                                  mean2 = women_empirical, std2 = np.sqrt(women_empirical*(1-women_empirical)), nobs2 = len(women_rank),
                                      alternative = 'two-sided')[1]

# Leave these lines here for grading and ease of debugging
print("Q5.1: p-value: ", q5_1)
print("Q5.2: p-value: ", q5_2)
print("Q5.3: p-value: ", q5_3)
```

Output:
```
Q5.1: p-value:  0.06606506021398857
Q5.2: p-value:  0.008010296971218134
Q5.3: p-value:  0.1717837459146108
```

### Answer 5

5.1) I found a two-sided p-value of around $0.07$. This is not statistically significant at the $0.05$ level because the p-value is greater than $0.05$, so we fail to reject the null hypothesis. Since we fail to reject the null hypothesis, we can not conclude that mean ratings of females and males are not equal. The underlying assumptions made using the t-test is that the samples are independent, and the variances are identical in both data sets. 

5.2) I found a two-sided p-value of around $0.008$. This is statistically significant at the $0.05$ level because the p-value is less than $0.05$, so we can reject the null hypothesis. Since we reject the null hypothesis, we can conclude that the fraction of men who would rate Star Wars a $5$ is not $51\%$.

5.3) I found a two-sided p-value of around $0.17$. This is not statistically significant at the $0.05$ level because the p-value is greater than $0.05$, so we fail to reject the null hypothesis. Since we fail to reject the null hypothesis, we can not conclude that women and men were equally likely to rate Star Wars a $5$.

5.4) 151 women rated Star Wars and 432 men rated Star Wars. There is a higher chance of a p-value being statistically significant with a larger sample size because a smaller p-value represents there is stronger evidence against the null hypothesis. In this case, about three times more men than women ranked Star Wars, so there is a larger sample size, which means that the null hypothesis is more likely to be rejected due to a smaller p-value. Since there are less women who rated Star Wars, the sample size is smaller meaning there is not as strong evidence against the null hypothesis, resulting in a larger p-value in which the null hypothesis fails to be rejected.

<!-- END QUESTION -->

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above will lose points.**

Tests for q1 (points: [1, 1, 2])
```
>>> isinstance(q1, tuple) and\
... (isinstance(q1[0], np.float64) or isinstance(q1[0], float)) and\
... (isinstance(q1[1], np.float64) or isinstance(q1[1], float))
True
```
```
>>> (0 < q1[0] < 100) and (0 < q1[1] < 100)
True
```
```
>>> from auto_checks import check_q1
>>> check_q1(q1)
'PASS'
```


Tests for q2_code (points: [1, 1, 2])
```
>>> q2_setting in [1, 2]
True
```
```
>>> isinstance(q2, np.float64) or isinstance(q2, float)
True
```
```
>>> from auto_checks import check_q2
>>> check_q2(q2, q2_setting)
'PASS'
```


Tests for q3 (points: [1, 1, 2, 1, 1, 2])
```
>>> isinstance(q3_1, tuple) and\
... (isinstance(q3_1[0], np.float64) or isinstance(q3_1[0], float)) and\
... (isinstance(q3_1[1], np.float64) or isinstance(q3_1[1], float))
True
```
```
>>> (0 < q3_1[0] < 100) and (0 < q3_1[1] < 100)
True
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> isinstance(q3_2, tuple) and\
... (isinstance(q3_2[0], np.float64) or isinstance(q3_2[0], float)) and\
... (isinstance(q3_2[1], np.float64) or isinstance(q3_2[1], float))
True
```
```
>>> (0 < q3_2[0] < 100) and (0 < q3_2[1] < 100)
True
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 2, 1, 1, 2])
```
>>> isinstance(q4_1, tuple) and\
... (isinstance(q4_1[0], np.float64) or isinstance(q4_1[0], float)) and\
... (isinstance(q4_1[1], np.float64) or isinstance(q4_1[1], float))
True
```
```
>>> (0 < q4_1[0] < 100) and (0 < q4_1[1] < 100)
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> isinstance(q4_2, tuple) and\
... (isinstance(q4_2[0], np.float64) or isinstance(q4_2[0], float)) and\
... (isinstance(q4_2[1], np.float64) or isinstance(q4_2[1], float))
True
```
```
>>> (0 <= q4_2[0] <= 5) and (0 <= q4_2[1] <= 5)
True
```
```
>>> from auto_checks import check_q4_2
>>> check_q4_2(q4_2)
'PASS'
```


Tests for q5 (points: [1, 1, 1, 1, 1, 1])
```
>>> import numpy
>>> isinstance(q5_1, numpy.float64) or isinstance(q5_1, float)
True
```
```
>>> from auto_checks import check_q5_1
>>> check_q5_1(q5_1)
'PASS'
```
```
>>> import numpy
>>> isinstance(q5_2, numpy.float64) or isinstance(q5_2, float)
True
```
```
>>> from auto_checks import check_q5_2
>>> check_q5_2(q5_2)
'PASS'
```
```
>>> import numpy
>>> isinstance(q5_3, numpy.float64) or isinstance(q5_3, float)
True
```
```
>>> from auto_checks import check_q5_3
>>> check_q5_3(q5_3)
'PASS'
```
